---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
free_preview: true
---

## test

```yaml
type: TabExercise
key: dffd96819c
xp: 100
```

The intention of this worksheet is to introduce you to some of the commands in the dplyr package in R (which we will continue to look at next week) and further familiarize you with the ggplot2 package.

As always, be sure to save your work as you progress through this file! To do so, use "File -> Save As" in the toolbar above and save it to the folder that you created for DATA100 in Week 1.  

Once saved, go to "Session -> Set Working Directory -> Choose Directory..." in the toolbar at the top of this page and navigate to your DATA100 folder.


**SETUP**

Before we get started, load the tidyverse package which contains the dplyr package.

`@pre_exercise_code`
```{r}

```

***

```yaml
type: NormalExercise
key: 1285c370b0
xp: 100
```

`@instructions`
test test test test test

`@hint`
nah

`@sample_code`
```{r}
install.packages("tidyverse")
library(tidyverse)
```

`@solution`
```{r}
install.packages("tidyverse")
library(tidyverse)
```

`@sct`
```{r}
ex() %>% check_library("tidyverse")
```
